# Workshop 8: Linear Regression


In this exercise, you will apply linear regression and Lasso regression methods to the dataset supplied to you, and then compare their results to determine whether Lasso regression is needed for this dataset:

**Dataset description**: You are provided a dataset with 20 variables. Variables $x1\ -\ x19$ refer to the independent variables, while variable $y$ is your dependent variable. Training data is stored in the file `/etc/data/regression-train.csv`.

**Note**: The TA will use a test set to verify your solution. The format (independent variables $x1\ -\ x19$, dependent variable  $y$) will be same, but TA's file may contain different number of data points than the split version from training set. Please ensure you take this into account, and do not hard code any dimensions.

---

```python
#Import necessary libraries
import pandas as pd
import numpy as np
import warnings
warnings.filterwarnings('ignore')
```

```python
#Read the data
df = pd.read_csv("./data/regression-train.csv")
X = df.iloc[:,:-1]
y = df.iloc[:,-1]
df
```

```python
from sklearn.model_selection import train_test_split
X_train_unscaled, X_test_unscaled, y_train, y_test = train_test_split(X,y,test_size=0.2,random_state=0)
```

```python
# For regression, it is particularly important to normalize our data before
# training the model, so we can better interpret our coefficients
from sklearn.preprocessing import StandardScaler
scaler = StandardScaler()
scaler.fit(X_train_unscaled)
X_train = scaler.transform(X_train_unscaled)
X_test = scaler.transform(X_test_unscaled)
#print(X_train[0])
```

We have given you a function to caculate the RMSE of a regression model, given its predictions and the ground truth y-values of the test data.

```python
import math
def calculate_rmse(y_true, y_pred):
  
  # Inputs:
  # y_true: ground truth dependent variable values, of type vector
  # y_pred: prediction outcomes from any regression method, with the same length as y_true
  
  # Outputs:
  # a single value of type double, with the RMSE value
    return(math.sqrt(np.sum((y_true - y_pred)**2)/len(y_true)))
```

## Part 1: Linear Regression (Group)
You will write code in the function `alda_regression_linear` to train simple linear regression. Detailed instructions for implementation and allowed packages have been provided the comments.

Before your begin, read the documentation on sklearn's [LinearRegression](https://scikit-learn.org/stable/modules/generated/sklearn.linear_model.LinearRegression.html).

<!-- grade_id: alda_regression_linear, points: 0, grade: false, solution: true -->
```python
from sklearn.linear_model import LinearRegression

def alda_regression_linear(X_train, X_test, y_train):
    # Perform linear regression
    # Inputs:
    # X_train: training data frame(19 variables, x1-x19)
    # X_test: test data frame(19 variables, x1-x19)
    # y_train: dependent variable, training data (vector, continous type)

    # Output:
    # A tuple containing:
    # - The regression model and 
    # - The list of predictions on test data (X_test) (vector) 
  
    # allowed packages: sklearn.linear_model
  
    # Function hints: Read the documentation for the functions LinearRegression (link above)
    
    # write code for building a linear regression model using X_train, y_train
    raise NotImplementedError()
```

```python
# Now test your model
lr_model, lr_predictions = alda_regression_linear(X_train, X_test, y_train)

print(f'Training RMSE: {calculate_rmse(y_train, lr_model.predict(X_train))}')
print(f'Test RMSE: {calculate_rmse(y_test, lr_predictions)}')
print('')

# Which attributes are most predictive of the outcome variable?
print(f'Model coefficients:\n{lr_model.coef_}')

```

<!-- grade_id: alda_regression_linear-public, points: 3, grade: true, solution: false -->
```python
lr_model, simple_linear_regression_result = alda_regression_linear(X_train, X_test, y_train)
np.testing.assert_equal(simple_linear_regression_result.shape, (27,))
np.testing.assert_almost_equal(calculate_rmse(y_test,simple_linear_regression_result),787.9099436300563)
np.testing.assert_almost_equal(lr_model.coef_[0], 48.18082989)
```

<!-- grade_id: alda_regression_linear-private, points: 2, grade: true, solution: false -->
```python
# Here's some test cases to make sure you're right
df_test = pd.read_csv("./data/regression-test.csv")
X_test2 = df_test.iloc[:,:-1]
y_test2 = df_test.iloc[:,-1]

lr_model, simple_linear_regression_result = alda_regression_linear(X, X_test2, y)
np.testing.assert_equal(simple_linear_regression_result.shape, (66,))
np.testing.assert_almost_equal(calculate_rmse(y_test2,simple_linear_regression_result),946.4318403560575)
# END HIDDEN TESTS
```

## Part 2: Lasso Regression (Group)
You will write code in the function `alda_regression_lasso` to train simple lasso regression. Detailed instructions for implementation and allowed packages have been provided the comments. 

Before your begin, read the documentation on sklearn's [LassoCV](https://scikit-learn.org/stable/modules/generated/sklearn.linear_model.LassoCV.html) - a Lasso regression model that uses CV to tune its hyperparameters.

**Note** that the lasso regression model has *built-in* crossvalidation, which it performs on the training dataset provided, to select the best shrinkage coefficient for the validation data.

<!-- grade_id: alda_regression_lasso, points: 0, grade: false, solution: true -->
```python
from sklearn.linear_model import LassoCV
def alda_regression_lasso(X_train, X_test, y_train, random_state=0):
    # Perform lasso regression
    # Inputs:
    # X_train: training data frame(19 variables, x1-x19)
    # X_test: test data frame(19 variables, x1-x19)
    # y_train: dependent variable, training data (vector, continous type)
    # random_state: a random state to use in CV model training
    # General Information:
        # use 10-fold cross validation to determine the best model hyperparameters
    
    # Output:
    # A tuple containing:
    # - The regression model and 
    # - The list of predictions on test data (X_test) (vector) 
  
    # allowed packages: sklearn.linear_model
  
    # Function hints: Read the documentation for the functions LassoCV (link above)
    
    # write code for lasso regression here
    # 10 fold cross validation
    # set up the random_state as 0 in order for reproducibility
    raise NotImplementedError()
```

**Before testing your model**: Do you expect the training error to be higher or lower? What about the testing error? What do you expect to be different about the coefficients?

```python
lasso_model, lasso_predictions = alda_regression_lasso(X_train, X_test, y_train)

# Should be ~541.7
print(f'Training RMSE: {calculate_rmse(y_train, lasso_model.predict(X_train))}')
# Should be ~639.4
print(f'Test RMSE: {calculate_rmse(y_test, lasso_predictions)}')
print('')

# Which attributes are most predictive of the outcome variable?
print(f'Model coefficients:\n{lasso_model.coef_}')

print()
# Note we called this 'lamda' in class, but sklearn calls it alpha (should be ~3.196)
print(f'The shinkage coefficient hyperparameter chosen by CV: {lasso_model.alpha_}')
```

<!-- grade_id: alda_regression_lasso-public, points: 3, grade: true, solution: false -->
```python
lasso_model, lasso_regression_result = alda_regression_lasso(X_train, X_test, y_train)
np.testing.assert_equal(lasso_regression_result.shape, (27,))
np.testing.assert_almost_equal(calculate_rmse(y_test, lasso_regression_result), 639.448499767289)
np.testing.assert_almost_equal(lasso_model.coef_[0], -201.29823887)
```

<!-- grade_id: alda_regression_lasso-private, points: 2, grade: true, solution: false -->
```python
# Here's some more test cases
df_test = pd.read_csv("./data/regression-test.csv")
X_test2 = df_test.iloc[:,:-1]
y_test2 = df_test.iloc[:,-1]

lasso_model, lasso_regression_result = alda_regression_lasso(X, X_test2, y)
np.testing.assert_equal(lasso_regression_result.shape, (66,))
np.testing.assert_almost_equal(calculate_rmse(y_test2, lasso_regression_result), 978.5702064182492)
```

From the results, compare the two regression models, including the training and testing RMSE, and the coefficients. Use the output of these functions to answer the following questions below:

1. The dataset contains 19 attributes. Are all 19 attributes useful for predicting the dependent variable? Why or why not? Use your results to justify the answer.
2. If not all attributes are predictive, use your Lasso model to perform feature selection. Which attributes should be kept? Use a correlation and/or scatter plot to justify your answer for at least one attribute (in a new cell below).

<!-- grade_id: comparison, points: 5, grade: true, solution: true -->
